Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Bellman-Ford
## Implementierung (15 Pkt.)
Implementieren Sie den Bellman-Ford Algorithmus, berechnen Sie für den gegebenen Graphen die Distanzmatrix.

![Graph](graph.svg)

Hinweis: Benutzen sie für Unendlich `inf` aus dem Paket `math`

<!-- grade_id: cell-7e09d49eb88dfeb7, points: 0, grade: false, solution: true -->
```python
from math import inf

# Adjacencylist
G = [[0,1,6],[0,2,5],[0,3,5],[3,2,-2],[2,1,-2],[1,4,-1],[2,4,1],[3,5,-1],[5,6,3],[4,6,3]];

def DistanceMatrix(graph):
    """ Call Bellman-Ford n times.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def BellmanFord(s, graph):
    """ Computes the shortest distance from s to every other node.
    
    Return: A list of distances.
    """
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-72a87cdf561fc897, points: 15, grade: true, solution: false -->
```python
G = [[0,1,6],[0,2,5],[0,3,5],[3,2,-2],[2,1,-2],[1,4,-1],[2,4,1],[3,5,-1],[5,6,3],[4,6,3]];
assert BellmanFord(0, G) == [0, 1, 3, 5, 0, 4, 3]
```

## Negative Kreise (10 Pkt.)
Modifizieren Sie den soeben implementierten Algorithmus so, dass negative Kreise erkannt werden und in diesem Fall eine Ausnahme mit entsprechender Nachricht erzeugt wird (`ValueError`).

<!-- grade_id: cell-dec1be90cc9f72fb, points: 10, grade: true, solution: false -->
```python
from unittest import TestCase

g = [[0,3,5],[0,1,4],[2,1,-10],[3,2,3],[1,3,5]]

t = TestCase()
with t.assertRaises(ValueError):
    BellmanFord(0,g)
```

Für den folgenden Graphen, nach wie vielen Iterationen im Bellman-Ford-Algorithmus kann man sich sicher sein, dass es einen negativen Kreis gibt?

![G2](graph2.svg)

<!-- grade_id: cell-1dcb77b55413701d, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Transitive Hülle (10 Pkt.)
In einem sozialen Netzwerk folgt Ron Steve, Steve Trish, Trish Ursula und Ursula Vivian:

![Soziales Netzverw](social_network.svg)

Wir betrachten dieses soziale Netzwerk als gerichteten Graphen wobei die relation **A folgt B** eine gerichtet Kante von A nach B bedeutet. Sei $E$ die Menge der Kanten und $E'$ die Menge der Kanten der transitiven Hülle, geben sie alle Kanten $\left(u,v\right) \in E'\setminus E$ an. 

<!-- grade_id: cell-4b4bd1b6eb5badff, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE